---
title: 'Covariance and Correlation'
description: 'Need to know the difference between the two, how to calculate the two (equations can help), advantages and disadvantages, when to use it and what for. Connection to causation.'
---

## Calculate covariance

```yaml
type: NormalExercise
key: 48671e7493
xp: 100
```

Typically, a data frame contains multiple variables. It is useful to know 

`@instructions`
<!-- Guidelines for instructions https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->
- Instruction 1
- Instruction 2

`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
```
