---
title: 'Multivariate Distributions'
description: 'Multivariate Distributions'
---

## Looking at Associations with Multivariate Categorical Data

```yaml
type: VideoExercise
key: af765926ee
xp: 50
```

`@projector_key`
0d3a97b8b574b58fe0baa4ecacff40e2

---

## Looking at Associations with Multivariate Quantitative Data

```yaml
type: VideoExercise
key: 56a52f2a07
xp: 50
```

`@projector_key`
da1d1e6b9320eeaa05d76da683f94c80

---

## Hands On Session-Multivariate Data

```yaml
type: VideoExercise
key: 15a65b2cfb
xp: 50
```

`@projector_key`
90cb0977a90a3bc7a7ad9dede63478db

---

## Calculate The Mean and Var

```yaml
type: NormalExercise
key: d72cbd39e4
xp: 100
```

<!-- Guidelines for contexts: https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->

`@instructions`
<!-- Guidelines for instructions https://instructor-support.datacamp.com/en/articles/2375526-course-coding-exercises. -->
- You have a pre loaded data frame with the data set: Weight Height named: df
- Calculate without using the method: mean, var the varience of the height of

`@hint`
<!-- Examples of good hints: https://instructor-support.datacamp.com/en/articles/2379164-hints-best-practices. -->
- This is an example hint.
- This is an example hint.

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}
# Examples of good success messages: https://instructor-support.datacamp.com/en/articles/2299773-exercise-success-messages.
```
